---
title: 'Model Evaluation using Confusion Matrix, Pseudo R-Square and AUC'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: a280b02057
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
