---
title: 'Introduction to Binary Logistic Regression'
description: 'Chapter description goes here.'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: f39c0d41c0
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
